```python
# Initialize Otter
import otter
grader = otter.Notebook("A4.ipynb")
```

# Assignment 4 (10 Points in total, will be 10% of your grade)

## How to complete assignments

There are both code and math/text components to these assignments.

### Math/text responses

Whenever you see:

```markdown
_Type your answer here, replacing this text._ 
```

You should editing this Markdown cell and insert your answers. 

If the question is asking for math use LaTeX. We understand not everyone is fluent in LaTeX but we think this is a good skill to learn. [Here is a fantastic tutorial from CalTech about using $\LaTeX$ in Jupyter Notebook.](http://chebe163.caltech.edu/2018w/handouts/intro_to_latex.html). You could also find various $\LaTeX$ tutorials and cheat sheets online. 

Some things to keep in mind:  Jupyter's Markdown-LaTeX will display things that don't work when you turn in the assignment. Here are some generic ideas that we find prevent trouble down the road. 
1. Test your output to see if it works well when you submit.  So execute the otter grader export and check the PDF to make sure it looks right!!! 
1. Avoid ```$x=1$ ``` use ```$$x=1$$``` instead
1.  No trailing space after the first dollar sign and before the second dollar sign. i.e. don't do `$ $ x = 5 $ $` use `$$x=5$$` instead
1. Markdown lists in the same cell after a math environment don't seem to work, and we don't know why
1. Avoid ```\align``` use ```\aligned``` instead

An example:


#### Question: Given a triangle with sides $x,y,z$ where $x=3$, $y=4$ calculate $z$

Your answer in the markdown should be

```
$$z = \sqrt{x^2 + y^2} = \sqrt{3^2+4^2} = \sqrt{25} = 5$$
```
which will render as 

$$z = \sqrt{x^2 + y^2} = \sqrt{3^2+4^2} = \sqrt{25} = 5$$


These responses will mostly be manually graded. There is leeway for individual style in notation, so don't worry too much as long as you don't write anything that is false. 

### Code responses

Whenever you see:

```python
answers = ...
```

You need to replace this section with some code that answers the questions and meets the specified criteria. Make sure you remove the 'raise' line when you do this (or your notebook will raise an error, regardless of any other code, and thus fail the grading tests).

You should write the answer to the questions in those cells (the ones with `...`), but you can also add extra cells to explore / investigate things if you need / want to. 

Any cell with `grader.check("Question_id")` statements in it is a test cell. You should not try to change or delete these cells. Note that there might be more than one assert that tests a particular question. Once you run the cell, you will see the public tests we provided for you.

If a test does fail, reading the error that is printed out should let you know which test failed, which may be useful for fixing it.

Note that some cells, including the test cells, may be read only, which means they won't let you edit them. If you cannot edit a cell - that is normal, and you shouldn't need to edit that cell.

## Question 1: Multiple Choice (1 Points)

For each questions, assign your answer to the list with the correct choices.

**Question 1.1:** Which of the following is a true property of the Kullback-Leibler divergence?

A): KL divergence is always negative.

B): KL divergence can be used to measure the distance between two distributions in both directions interchangeably, i.e. $KL(P||Q) = KL (Q||P)$ is true for every distribution $Q$ and $P$

C): KL divergence is zero if and only if the two distributions being compared are identical.

D): KL divergence satisfies the triangle inequality.

**Question 1.2:** When using the KL divergence to compare two probability distributions $P$ and $Q$, a small KL divergence value indicates that:

A): The distributions $P$ and $Q$ are very similar.

B): The distribution $Q$ is a good approximation of distribution $P$.

C): We would not be very surprised when observing data generated by distribution 
$P$ when our model of the data was $Q$.

D): All of the above

**Question 1.3:** The perplexity parameter in t-SNE:

A): Is independent of the dataset size and only affects the computational complexity of the algorithm.

B): Directly controls the learning rate of the gradient descent optimization in t-SNE.

C): Can be thought of as a measure of the effective number of neighbors for each point and influences the balance between local and global aspects of the data.

D): Sets the number of dimensions to which the data will be reduced.

**Question 1.4:** Which of the following best describes the principle of operation for both t-SNE and UMAP?

A): They both linearly project data into lower-dimensional spaces while preserving global data structure.

B): They both aim to preserve the local neighborhood structure of the data in a lower-dimensional space.

C): They both primarily focus on preserving the pairwise distances between all points in the dataset.

D): They both use a global optimization technique to ensure that similar data points in the high-dimensional space are mapped far apart in the lower-dimensional representation.

**Question 1.5:** The absolute positions of clusters within the embedding space have a direct and interpretable meaning.

A): True

B): False

**Question 1.6** Compared to t-SNE, UMAP:

A) Is generally faster and more scalable, particularly on larger datasets.

B) Can only visualize datasets in two or three dimensions, while t-SNE can project data into any number of dimensions.

C) Lacks the ability to preserve local neighborhood structures of the data.

D) Requires more hyperparameters to be tuned for optimal performance.

_Points:_ 1

```python
multiple_choices = ...
```

```python
grader.check("11_multiple_choice")
```

## Question 2: KL-Divergence (2 Points)

The KL-Divergence is a measure of how one probability distribution $p(x)$ is different from another distribution $q(x)$.

$$\begin{aligned}
D_{KL}(p(x)||q(x)) &= \mathbb{E}_p(x)\left[\log\frac{p(x)}{q(x)}\right] \\
&= \int_{-\infty}^{\infty} p(x)\log \frac{p(x)}{q(x)} dx & \text{for continuous distribution $p$ and $q$} \\
&= \sum p(x)\log \frac{p(x)}{q(x)} & \text{for discrete distribution $p$ and $q$} \\
\end{aligned}$$

<!-- BEGIN QUESTION -->

### Question 2.1: Prove that KL Divergence is Non-negative

Prove that 
$$D_{KL}(p(x)||q(x))\geq 0$$

Hint:
- Apply Jensen's inequality, which says that $\mathbb{E}[f(z)]\geq f(\mathbb{E}[z])$ for any convex function $f:\mathbb{R}\to\mathbb{R}$ and for any random variable $z$.
- In this case, you should use $f(z) = -\log(z)$
- Note that $\log\frac{a}{b} = \log(a)-\log(b) = -(\log(b)-\log(a)) = -\log\frac{b}{a}$
- Think about the case where the divergence would be lowest... when the two distributions are the same!

_Points:_ 2

_Type your answer here, replacing this text._

<!-- END QUESTION -->

## Question 3: Explore PCA v.s. TSNE v.s. UMAP (4 Points)

Let's have you use some interesting data to explore the differences and similarities between dimensionality reduction techniques.

First off we will get UMAP installed for you. 

Then we will download some data with stock prices for 56 large global companies between 2003 and 2007 (a time when the market overall was growing rapidly and before the 2008 crash). These companies come from several different kinds of industries... tech, finance, manufacturing, pharmaceutical, etc.  If you aren't already familiar with these companies and their industries, you really should investigate a little or you won't be able to talk about the results of your analysis properly.  

Lastly we will transform the data in ways that make it easier to find patterns... please make sure you understand how and why these transformations are helpful!

```python
# first off we will need umap
# this is maybe not necessary if you have the newest sklearn
# but assuming you have our environment for this class it is needed
%pip install umap-learn
```

```python
import sys

import numpy as np
import pandas as pd


import seaborn as sns
import matplotlib.pyplot as plt
%matplotlib inline
%config InlineBackend.figure_format='retina'
sns.set(style='white', context='notebook', rc={'figure.figsize':(12,8)})

# plotting helper function
def plot_components(components,x='pc1',y='pc2'):
    '''
    2d scatter plot of dimensionality reductions with labelling

    inputs:
        components (pd.DataFrame) shape is (samples x components), index is sample names
        x (str) name of column in components to plot as x axis
        y (str) name of column in components to plot as y axis

    typically you might generate components like this
        pcs = PCA(n_components=2).fit_transform(daily_change)
        components = pd.DataFrame(pcs,columns=['pc1','pc2'])
        components.index = daily_change.index
    '''
    sns.scatterplot(data=components, x=x, y=y)
    sns.despine()
    ax = plt.gca()
    for company in components.index:
      ax.annotate(company, xy=components.loc[company])

symbol_dict = {
    "TOT": "Total",
    "XOM": "Exxon",
    "CVX": "Chevron",
    "COP": "ConocoPhillips",
    "VLO": "Valero Energy",
    "MSFT": "Microsoft",
    "IBM": "IBM",
    "TWX": "Time Warner",
    "CMCSA": "Comcast",
    "CVC": "Cablevision",
    "YHOO": "Yahoo",
    "DELL": "Dell",
    "HPQ": "HP",
    "AMZN": "Amazon",
    "TM": "Toyota",
    "CAJ": "Canon",
    "SNE": "Sony",
    "F": "Ford",
    "HMC": "Honda",
    "NAV": "Navistar",
    "NOC": "Northrop Grumman",
    "BA": "Boeing",
    "KO": "Coca-Cola",
    "MMM": "3M",
    "MCD": "McDonald's",
    "PEP": "Pepsi",
    "K": "Kellogg",
    "UN": "Unilever",
    "MAR": "Marriott",
    "PG": "Procter Gamble",
    "CL": "Colgate-Palmolive",
    "GE": "General Electric",
    "WFC": "Wells Fargo",
    "JPM": "JPMorgan Chase",
    "AIG": "AIG",
    "AXP": "American Express",
    "BAC": "Bank of America",
    "GS": "Goldman Sachs",
    "AAPL": "Apple",
    "SAP": "SAP",
    "CSCO": "Cisco",
    "TXN": "Texas Instruments",
    "XRX": "Xerox",
    "WMT": "Wal-Mart",
    "HD": "Home Depot",
    "GSK": "GlaxoSmithKline",
    "PFE": "Pfizer",
    "SNY": "Sanofi-Aventis",
    "NVS": "Novartis",
    "KMB": "Kimberly-Clark",
    "R": "Ryder",
    "GD": "General Dynamics",
    "RTN": "Raytheon",
    "CVS": "CVS",
    "CAT": "Caterpillar",
    "DD": "DuPont",
}


symbols, names = np.array(sorted(symbol_dict.items())).T

quotes = []

for symbol in symbols:
    print("Fetching quote history for %r" % symbol, file=sys.stderr)
    url = (
        "https://raw.githubusercontent.com/scikit-learn/examples-data/"
        "master/financial-data/{}.csv"
    )
    quotes.append(pd.read_csv(url.format(symbol)))

close_prices = np.vstack([q["close"] for q in quotes])
open_prices = np.vstack([q["open"] for q in quotes])

# Stock prices are non-stationary time series data... 
# if you don't know what that means read this https://otexts.com/fpp2/stationarity.html
# taking the daily difference in stock price helps with non-stationarity by removing big trends
# looking at that daily difference as a ratio of change
# normalizes somewhat across stocks with very different prices
# and may help us discover groups of stocks that 
# move up and down in lock step but at different raw prices
daily_diff = close_prices - open_prices
daily_change = pd.DataFrame(daily_diff/open_prices, index=names, columns=quotes[0]['date'])
daily_change.index.name = 'company'
daily_change
```

### Question 3.1: PCA plot

Load the necessary things from scikit-learn, do the data transform, use the helper function to plot the results

_Points:_ 0.25

```python
from sklearn.decomposition import PCA

pca = ...

# call fit_transform on pca
pcs = ...
# transform it into data frame
pcs = pd.DataFrame(pcs,columns=['pc1','pc2'])
pcs.index = daily_change.index
# plot the components
plot_components(pcs)
```

```python
grader.check("31_plot_pca")
```

### Question 3.2: t-SNE plot

Load the necessary things from scikit-learn, do the data transform, use the helper function to plot the results

Use a `random_state=99` and explore the following values for `perplexity`: 5, 20, and 40 (and any others you feel like)

_Points:_ 0.25

```python
...

tsne_perp_5 = ...
tcs = ...
# construct the data frame
tcs = pd.DataFrame(tcs,columns=['c1','c2'])
tcs.index = daily_change.index

plot_components(tcs,'c1','c2')
```

```python
tsne_perp_20 = ...
tcs = ...
tcs = pd.DataFrame(tcs,columns=['c1','c2'])
tcs.index = daily_change.index

plot_components(tcs,'c1','c2')
```

```python
tsne_perp_40 = ...
tcs = ...
tcs = pd.DataFrame(tcs,columns=['c1','c2'])
tcs.index = daily_change.index

plot_components(tcs,'c1','c2')
```

```python
grader.check("32_plot_tsne")
```

### Question 3.3: UMAP plot

Load the necessary things from umap, do the data transform, use the helper function to plot the results

Use a `random_state=99` and `min_dist=0.1` and explore the following values for `n_neighbors`: 5, 20, and 40 (and any others you feel like)

You'll probably want to take a look at the UMAP docs: https://umap-learn.readthedocs.io/en/latest/parameters.html



_Points:_ 0.25

```python
from umap import UMAP
```

```python
umap_5_neigh = ...
ucs = ...

ucs = pd.DataFrame(ucs,columns=['c1','c2'])
ucs.index = daily_change.index
plot_components(ucs,'c1','c2')
```

```python
umap_20_neigh = ...
ucs = ...

ucs = pd.DataFrame(ucs,columns=['c1','c2'])
ucs.index = daily_change.index
plot_components(ucs,'c1','c2')
```

```python
umap_40_neigh = ...
ucs = ...

ucs = pd.DataFrame(ucs,columns=['c1','c2'])
ucs.index = daily_change.index
plot_components(ucs,'c1','c2')
```

```python
grader.check("33_plot_umap")
```

<!-- BEGIN QUESTION -->

### Question 3.4: Compare and discuss

Let's discuss what you saw with the various plots. 

Particularly we are looking to see how companies appear to be similar to each other in the reduced dimensions.  If two companies are near each other in low-d assumedly their stocks move up and down together. This might happen because they are in the same industry and there are good weeks and bad weeks for oil stocks that are different than good/bad weeks for tech stocks.  But industry isn't everything; clearly inside an industry there could be companies doing well while others do poorly. So looking for clusters of industry is simply something I assume might be there, but doesn't have to be. Think of this as the default "pattern" I'm asking you about in the questions below.

**NOTE:** I'm not a broker or an economist, so there are probably other reasons for companies clustering together or separating from each other. If you have particular expertise or are gung-ho to do some research about these companies in this time period I'm curious to see what other things you come up with. So hit me with your best explanations for why "patterns" exist in the low-d data! Think of this as a mini project... this is the kind of subject matter expertise that you will need to bring to your projects :)

So in the Markdown box below please address the following questions in a short essay format. _Your response should probably be 1 (and no more than 2!) paragraphs per bullet point below._ 
- Characterize what the PCA plot looks like, say what patterns you can see
- Characterize differences among t-SNE plots, say what patterns you can see. Talk about how these patterns appear/disappear/modify with the different parameters. Relate these pattern changes to what you know about the nature of high and low perplexity
- Characterize differences among UMAP plots, say what patterns you can see. Talk about how these patterns appear/disappear/modify with the different parameters. Relate these pattern changes to what you know about the nature of high and low n_neighbors
- Why didn't I ask you to investigate even larger values for perplexity or n_neighbors, like e.g. 100? What happens when you try that with the algorithm?
- What differences do you see in the plots between the 3 algorithms? Do you feel like there's a strong reason to prefer one of these algorithms? Reasons might include results, computational complexity, hyper-parameter tuning, and anything else you feel is relevant
- Given what you've seen here, what conclusions do you feel comfortable drawing about these stocks and why? What else do you suspect but would feel like you can't prove given this analysis?


_Points:_ 3.25

_Type your answer here, replacing this text._

<!-- END QUESTION -->

## Question 4: Use UMAP as the Basis for a Supervised Classifier (4 Points)

Heres an example of how to use UMAP as the input to a supervised learning classifier... You'll have to look some stuff up about UMAP https://umap-learn.readthedocs.io/en/latest/transform.html and k-Nearest-Neighbors https://scikit-learn.org/stable/modules/neighbors.html#nearest-neighbors-classification

In this example, we will:

- get the full 70k sample x 784 dimension MNIST dataset
- create a pipeline to 
    - make a dimension reduction using UMAP down to a 2 dimension embedding
    - use a k-Nearest Neighbors classifier in that 2d embedding to classify digits
- visualize the results of the predictions in the embedding 

Some interesting notes on this process:

- UMAP and kNN are similar in that large numbers of neighbors produces smoother results.. more global relationships for UMAP and more smooth decision boundaries for kNN. Both of these are _VERY_ dependent on hyperparameters!
- We are setting you up with a pipeline version of this process in order to facilitate you doing some hyperparameter selection in the next question.

```python
from sklearn.metrics import classification_report
from sklearn.model_selection import train_test_split, GridSearchCV
from sklearn.neighbors import KNeighborsClassifier
from sklearn.pipeline import Pipeline
from sklearn.datasets import fetch_openml

from umap import UMAP

import numpy as np
import pandas as pd

import matplotlib.pyplot as plt
%matplotlib inline
%config InlineBackend.figure_format='retina'
import seaborn as sns
sns.set_theme()

# fetch dataset
mnist = fetch_openml('mnist_784', version=1)
mnist.target = mnist.target.astype(int)
```

```python
# 50/50 train test split
X_train, X_test, y_train, y_test = train_test_split(
    mnist.data, mnist.target, test_size=0.5
)

# arbitrary parameters, feel free to play with this
trans = UMAP(n_neighbors=5, min_dist=0.1)
knn = KNeighborsClassifier(n_neighbors=5)

pipe = Pipeline([
    ('dim_reduce', trans),
    ('classifier', knn)
])

pipe.fit(X_train,y_train)
```

```python
# this is how you can extract the embedding from the training set data
# .steps is a list of the steps of the pipeline
# [0] is the first one (dimred)
# which yields the dimreduce tuple from the Pipeline command, 
# [1] is the actual umap class while [0] is just the name
# finally inside the umap instance .embedding_ is the transformed training data
embedding_train = pipe.steps[0][1].embedding_
```

```python
plt.figure(figsize=(10, 10))
scatter = plt.scatter(embedding_train[:, 0], embedding_train[:, 1], c=y_train.astype(int), s=0.1, cmap='Spectral')
handles, _ = scatter.legend_elements(prop="colors", alpha=0.6) # use my own labels
legend1 = plt.legend(handles, np.arange(0, 10))
plt.title('Test set, colored by prediction');
plt.title('Training set, colored by true label')
plt.xlabel("Embedding Dimension 1")
plt.ylabel("Embedding Dimension 2");
```

### Question 4.1: Make a prediction on the test set

_Points:_ 1

```python
# how can we now get predictions from the fitted pipeline? basic scikit-learn usage!
prediction = ...
print(classification_report(y_test, prediction))
```

```python
grader.check("41_UMAP_Basis")
```

```python
# to get the test set embedding we need must explictly call dimred.transform() on the test data
# to call dimred.transform() we use the pipe.steps trick again

# note this .transform() is an approximation!!!
# The vanilla UMAP cannot transform test data its never seen, because the embedding is produced
# during a gradient descent process. However there is various black magic being done here
# where another supervised algorithm is trained to try to go from original training data to 
# the low d embedding, and that supervised algorithm is used to transform the test data
# if you're interested here's discussions of HOW this black magic happens https://github.com/lmcinnes/umap/issues/40

embedding_test = pipe.steps[0][1].transform(X_test) 
```

```python
# look at all that salt and pepper in scatter dots in between the "clusters"
# these are "outlier" digits that don't really look like anything in the training data
# so they really don't fall where other digits of that label did
plt.figure(figsize=(10, 10))
scatter = plt.scatter(embedding_test[:, 0], embedding_test[:, 1], c=prediction.astype(int), s=0.1, cmap='Spectral');
handles, _ = scatter.legend_elements(prop="colors", alpha=0.6) # use my own labels
legend1 = plt.legend(handles, np.arange(0, 10))
plt.title('Test set, colored by prediction');
plt.xlabel("Embedding Dimension 1")
plt.ylabel("Embedding Dimension 2");
```

### Question 4.2: Model Selections

OK its your turn!!

As you can see a pipeline sets you up to do proper model selection on hyperparameters.

We had several notebooks that covered gridsearch (8, 12, & 15); the last two addressed it specifically with pipelines.  But this will be a little different as we are using a supervised learning algorithm on top of the dimensionality reduction, and the metrics you will probably want are things like accuracy rather than BIC.  Heres' a sklearn demo that uses supervised learning on top of dimensionality reduction that's relevant https://scikit-learn.org/stable/auto_examples/compose/plot_compare_reduction.html#sphx-glr-auto-examples-compose-plot-compare-reduction-py

Anyway... go for it! Pick a semi reasonable set of parameters for UMAP and kNN and see how good you can make the test set error!

**Notes:** `GridSearchCV` is very computationally intensive (think about why). Therefore, you need to choose your hyperparamter grid wisely. We advise that you do not exceed a total of 16 possible models for your grid search, but feel free to have more if you have a fancy multi-core CPU.

_Points:_ 2

```python
# We will continue to use the NMIST dataset for this question
pipe = Pipeline([
    # fill in your pipeline
    (...), # UMAP dimensionality reduction
    (...), # kNN classifier
])

...

param_grid = {
    "...": ..., # UMAP hyper-parameters
    "...": ..., # kNN hyper-parameters
}

...

# create a GridSearchCV Object with the pipe and the param_grid, you can set n_jobs=-1 to enable parallelism
# Do a two fold validation by setting the cv=2, and verbose=5
grid = ...
```

```python
# fit the grid with the training data
# this will take considerable amount of time.
# feel free to train on the whole dataset, but for the time being,
# you can train on the first 5k samples

fitted = grid.fit(X_train[:5000], y_train[:5000]) 
```

```python
grader.check("42_UMAP_model_selection")
```

<!-- BEGIN QUESTION -->

### Question 4.3: Model Performance Evaluations

The following code visualize the performance of each individual model in the `GridSearchCV`. Comment on which model performs the best, with which set of hyperparameter. 

```python
results = pd.DataFrame(grid.cv_results_['params'] )
results['score'] = grid.cv_results_['mean_test_score']
# heres a tibble of your grid search results
results.head()
```

```python
# this makes a 2d heatmap comparing JUST TWO of the parameters you are interested in 

# YOU MUST INTERVENE HERE
# the first two args to .pivot (index, columns) 
# must be the names of the parameters you want to score
# right now this assumes you are comparing the first two params
# modify it according to what you need to visualize!!!
# if you're doing 3 params, you may need to make 2 heatmaps, one for each pair of params

grid_params = results.columns
heatvals = results.pivot(index=grid_params[0],
                           columns=grid_params[1],
                           values='score')
plot = sns.heatmap(heatvals, annot=True, fmt='5.4f', robust=True, cmap="crest");
plot.figure.savefig("./data/43plot.png")
heatvals
```

```python
# NB some people have a problem where you only see the first row
# of the numeric annotations in the heatmap above
# if that affects you no big deal.  thats why I made it print the results grid too
# if you want, you can try to uncomment the following, run it, and restart your kernel
# %pip install seaborn --upgrade 
```

![](./data/43plot.png)

_Points:_ 1

_Type your answer here, replacing this text._

<!-- END QUESTION -->

## The End of A4

Have a look back over your answers, and also make sure to `Restart & Run All` from the kernel menu to double check that everything is working properly. This restarts everything and runs your code from top to bottom.

Once you're happy with your work, click the disk icon to save, and submit the zip file onto gradescope. **You MUST submit all the required component to receive credit.**

Note that you can submit at any time, but **we grade your most recent submission**. This means that **if you submit an updated notebook after the submission deadline, it will be marked as late**.

_Note:_ If you encountered `LatexFailed` message during exporting, or has the `AttributeError: module 'nbconvert' has no attribute 'pdf'` error, it indicates that your $\LaTeX$ code is not correct. Try to the LaTeX syntax error by scrolling up to see the LaTeX error message. If you need any additional helps, please make a private post on campuswire and we are happy to help.

## Submission

Make sure you have run all cells in your notebook in order before running the cell below, so that all images/graphs appear in the output. The cell below will generate a zip file for you to submit.

Please make sure to see the output of the gradescope autograder. You are responsible for waiting and ensuring that the autograder is executing normally for your submission. Please create a campuswire post if you see errors in autograder execution.

```python
grader.export(force_save=True, run_tests=True, files=['data'])
```

Tests for 11_multiple_choice (points: 1)
```
>>> assert isinstance(multiple_choices, list)
>>> assert all([choice in ['A', 'B', 'C', 'D'] for choice in multiple_choices])
```
```
>>> assert multiple_choices[0] == 'C'
```
```
>>> assert multiple_choices[1] == 'D'
```
```
>>> assert multiple_choices[2] == 'C'
```
```
>>> assert multiple_choices[3] == 'B'
```
```
>>> assert multiple_choices[4] == 'B'
```
```
>>> assert multiple_choices[5] == 'A'
```


Tests for 31_plot_pca (points: 0.25)
```
>>> assert isinstance(pca, PCA)
>>> assert isinstance(pcs, pd.DataFrame)
```
```
>>> assert np.allclose((2, 1258), pca.components_.shape)
```


Tests for 32_plot_tsne (points: 0.25)
```
>>> params = tsne_perp_5.get_params()
>>> assert np.isclose(params['perplexity'], 5)
>>> assert np.isclose(params['random_state'], 99)
```
```
>>> params = tsne_perp_20.get_params()
>>> assert np.isclose(params['perplexity'], 20)
>>> assert np.isclose(params['random_state'], 99)
```
```
>>> params = tsne_perp_40.get_params()
>>> assert np.isclose(params['perplexity'], 40)
>>> assert np.isclose(params['random_state'], 99)
```


Tests for 33_plot_umap (points: 0.25)
```
>>> params = umap_5_neigh.get_params()
>>> assert np.isclose(params['n_neighbors'], 5)
>>> assert np.isclose(params['random_state'], 99)
```
```
>>> params = umap_20_neigh.get_params()
>>> assert np.isclose(params['n_neighbors'], 20)
>>> assert np.isclose(params['random_state'], 99)
```
```
>>> params = umap_40_neigh.get_params()
>>> assert np.isclose(params['n_neighbors'], 40)
>>> assert np.isclose(params['random_state'], 99)
```


Tests for 41_UMAP_Basis (points: 1)
```
>>> assert isinstance(prediction, np.ndarray)
>>> assert np.isclose(prediction.shape, y_test.shape)
```


Tests for 42_UMAP_model_selection (points: 2)
```
>>> assert isinstance(pipe, Pipeline)
>>> assert isinstance(grid, GridSearchCV)
```
