```python
# Initialize Otter
import otter
grader = otter.Notebook("A1.ipynb")
```

# Assignment 1 (10 Points in total, will be 10% of your grade)

## How to complete assignments

There are both code and math/text components to these assignments.

### Math/text responses

Whenever you see:

```markdown
_Type your answer here, replacing this text._ 
```

You should editing this Markdown cell and insert your answers. 

If the question is asking for math use LaTeX. Tutorials on using LaTeX ( Scott insert something here).  Also some generic advice.  Avoid XXX . Use YYY instead. (Scott what was the ams-math environment that gave us problems last quarter??)

An example:


#### Question: Given a triangle with sides $x,y,z$ where $x=3$, $y=4$ calculate $z$

Your answer in the markdown should be

```
$z = \sqrt{x^2 + y^2} = \sqrt{3^2+4^2} = \sqrt{25} = 5$
```
which will render as 

$z = \sqrt{x^2 + y^2} = \sqrt{3^2+4^2} = \sqrt{25} = 5$

_Note:_ Be sure to not having any trailing space after the first dollar sign and before the second dollar sign. i.e. Don't do `$ x = 5 $`, instead, do `$x=5$`

These responses will mostly be manually graded. There is leeway for individual style in notation, so don't worry too much as long as you don't write anything that is false. 

### Code responses

Whenever you see:

```python
answers = ...
```

You need to replace this section with some code that answers the questions and meets the specified criteria. Make sure you remove the 'raise' line when you do this (or your notebook will raise an error, regardless of any other code, and thus fail the grading tests).

You should write the answer to the questions in those cells (the ones with `...`), but you can also add extra cells to explore / investigate things if you need / want to. 

Any cell with `grader.check("Question_id")` statements in it is a test cell. You should not try to change or delete these cells. Note that there might be more than one assert that tests a particular question. Once you run the cell, you will see the public tests we provided for you.

If a test does fail, reading the error that is printed out should let you know which test failed, which may be useful for fixing it.

Note that some cells, including the test cells, may be read only, which means they won't let you edit them. If you cannot edit a cell - that is normal, and you shouldn't need to edit that cell.

## Q1: Basic Probability (0.75 Points)

For two events $A$ and $B$, it is given that $P(A) = 0.73$, $P(B) = 0.48$ and $P(A \cap B) = 0.29$.

<!-- BEGIN QUESTION -->

### Q1.1 Find $P(A|B)$

_Points:_ 0.25

_Type your answer here, replacing this text._

<!-- END QUESTION -->

<!-- BEGIN QUESTION -->

### Q1.2 Find $P(B|A)$

_Points:_ 0.25

_Type your answer here, replacing this text._

<!-- END QUESTION -->

<!-- BEGIN QUESTION -->

### Q1.3 Determine whether or not $A$ and $B$ are independent.

_Points:_ 0.25

_Type your answer here, replacing this text._

<!-- END QUESTION -->

## Q2: Drawing Bills from Boxes (0.8 Points)

There are three boxes: a red box, a green box, and a blue box.

* The **red** box has one 100-dollar bill and nine 1-dollar bills.
* The **green** box has five 100-dollar bills and five 1-dollar bills
* The **blue** box has nine 100-dollar bills and one 1-dollar bill.

One box is chosen at random with the probabilities P(red) = 0.6, P(Green) = 0.3, and P(Blue) = 0.1. Then, a bill is picked at random from the chosen box.

<!-- BEGIN QUESTION -->

### Q2.1 What is the probability of picking a 100-dollar bill?

_Points:_ 0.2

_Type your answer here, replacing this text._

<!-- END QUESTION -->

<!-- BEGIN QUESTION -->

### Q2.2 What is the probability of picking a 1-dollar bill?

_Points:_ 0.2

_Type your answer here, replacing this text._

<!-- END QUESTION -->

<!-- BEGIN QUESTION -->

### Q2.3 Given that the picked bill is a 100-dollar bill, what is the probability that it came from the Green box?

_Points:_ 0.2

_Type your answer here, replacing this text._

<!-- END QUESTION -->

<!-- BEGIN QUESTION -->

### Q2.4 Let’s draw a random bill out of the box. What is the expected value of the dollar worth of the bill? What does that mean?

_Points:_ 0.2

_Type your answer here, replacing this text._

<!-- END QUESTION -->

## Q3 Distribution 101 (0.6 Point)

<!-- BEGIN QUESTION -->

### Q3.1 

![title](imgs/distribution.png)

Here is some data. Do you think this would be well-fit by a single normal distribution? If not, is there another distribution that you would suggest to fit this data? Do you think this data would be well fit by multiple normal distributions? If so, how many would you suggest, and why that? Don’t freak out here… while there are better and worse answers to this question there is not just a single right answer.  

_Points:_ 0.2

_Type your answer here, replacing this text._

<!-- END QUESTION -->

<!-- BEGIN QUESTION -->

### Q3.2
Picking a distribution by eye like we did above is NOT good.  It’s best if you have a theoretical reason, based on the math of the thing you are modeling. With that in mind, what kind of distribution (or distributions!) would you expect to use to model _____ and why:
1. Whether a coin is fair or not?
2. How frequently we expect customers to enter a store?
3. Height of male college basketball players
4. Height among all college freshmen


_Points:_ 0.4

_Type your answer here, replacing this text._

<!-- END QUESTION -->

## Q4 Expected Value and Variance (0.85 Points)
Prove the following properties of Bernoulli Distribution.

<!-- BEGIN QUESTION -->

### Q4.1

$$\sum_{x=0}^1 p(x|\mu) = 1$$

_Points:_ 0.25

_Type your answer here, replacing this text._

<!-- END QUESTION -->

<!-- BEGIN QUESTION -->

### Q4.2 

$$\mathbb{E}[x] = \mu$$

_Points:_ 0.25

_Type your answer here, replacing this text._

<!-- END QUESTION -->

<!-- BEGIN QUESTION -->

### Q4.3 

$$Var[x] = \mu(1-\mu)$$

_Points:_ 0.35

_Type your answer here, replacing this text._

<!-- END QUESTION -->

## Q5 Monty Hall (1 Points)
In the Monty Hall problem, without loss of generality, consider that you originally chose Door 1, Monty opened Door 2, and $P(\text{choose}=1, \text{open}=2) = a$, calculate the following probabilities:

<!-- BEGIN QUESTION -->

### Q5.1 
What is the probability of the car being behind Door 1 given that you chose Door 1 and Monty opened Door 2? i.e. $$P(\text{car}=1|\text{choose}=1, \text{open}=2)$$

_Points:_ 0.3

_Type your answer here, replacing this text._

<!-- END QUESTION -->

<!-- BEGIN QUESTION -->

### Q5.2 

What is the probability of the car being behind Door 3 given that you chose Door 1 and Monty opened Door 2?
$$P(\text{car}=3|\text{choose}=1, \text{open}=2)$$


_Points:_ 0.3

_Type your answer here, replacing this text._

<!-- END QUESTION -->

<!-- BEGIN QUESTION -->

### Q5.3 
Compare your answers from part(1) and part(2), which has a higher probability and should you switch the door? 

_Points:_ 0.4

_Type your answer here, replacing this text._

<!-- END QUESTION -->

## Q6: MLE (2 Points)

There are two datasets:

- Coin flips (0=tails,1=heads): $\mathcal{D}_1 = [0, 1, 1, 1, 1, 1, 0, 1, 1, 1]$
- Male heights (cm): $\mathcal{D}_2 = [202., 193., 190., 189., 210., 188., 175., 185., 152., 182.]$

For each dataset answer the following questions...


<!-- BEGIN QUESTION -->

### Q6.1 What is the correct distribution to describe each dataset AND WHY?


_Points:_ 0.4

_Type your answer here, replacing this text._

<!-- END QUESTION -->

<!-- BEGIN QUESTION -->

### Q6.2 
Write down the likelihood function $P(D|\theta)$ for each distribution. 

Note this will be in term of a product of PDFs for each datapoint in the dataset $D$ because these processes are assumed to be statistically independent.   


_Points:_ 0.6

_Type your answer here, replacing this text._

<!-- END QUESTION -->

<!-- BEGIN QUESTION -->

### Q6.3 Write down the log likelihood function only for the coin datasets, and simplify it enough to get rid of products and exponents.

_Note:_ You don't have to calculate the exact likelihood values, you can use $x_k$ to represent the $k^{\text{th}}$ value in the 

_Points:_ 0.4

_Type your answer here, replacing this text._

<!-- END QUESTION -->

<!-- BEGIN QUESTION -->

### Q6.4 Only for the coin dataset: how would you use the log likelihood function to analytically solve for the MLE of $\theta$? 
Note you don't have to do all the derivation and simplification, but kudos to you if you can.  At the very least describe the procedure in words and/or sketch out  the beginning of the derivation. The idea is we want you to demonstrate that you understand the concept of MLE, so whatever you think is sufficient to do that.



_Points:_ 0.4

_Type your answer here, replacing this text._

<!-- END QUESTION -->

<!-- BEGIN QUESTION -->

### Q6.5 For both datasets, what is the equation for the MLE of $\theta$?  What are the values for each dataset? 
Even if you didn't derive it from first principles, you should still know what the equation is for the MLE. Write down the proper equation, and calculate the value (actual number) of that MLE.

_Points:_ 0.2

_Type your answer here, replacing this text._

<!-- END QUESTION -->

## Q7 Bayesian Estimation (2 Points)

Let's use Bayesian techniques to tackle the same data on coin flips we previously addressed with MLE (0=tails,1=heads): $\mathcal{D}_1 = [0, 1, 1, 1, 1, 1, 0, 1, 1, 1]$

<!-- BEGIN QUESTION -->

### Q7.1 What is the correct conjugate prior?
Given the likelihood function you already wrote down in Q6.2, what's the proper conjugate prior distribution for that likelihood function? Why that one?


_Points:_ 0.3

_Type your answer here, replacing this text._

<!-- END QUESTION -->

<!-- BEGIN QUESTION -->

### Q7.2 Write down that Bayes equation for calculating the posterior distribution $P(\theta | D)$
Do this in terms of the likelihood fuction and prior you have selected.  Use $n$ as the number of flips and $x$ as the number of heads in those flips.

_Points:_ 0.3

_Type your answer here, replacing this text._

<!-- END QUESTION -->

<!-- BEGIN QUESTION -->

### Q7.3 Special interprestation of prior parameters
There is a special interpretation of the parameters for the prior distribution, where we can talk about them as a virtual coin flipping experiment. Describe the relationship in words. Tell us WHY this works in either words or math, your choice.

_Hint:_ see the slides from lecture, or if you simplified your answer to 7.2 above you can see it there

_Points:_ 0.2

_Type your answer here, replacing this text._

<!-- END QUESTION -->

### Q7.4 What parameters of the prior would make it  "uninformative" ? 

Plot the uniformative prior using the following stub, replacing `DISTRIBUTION` with a string containing the scipy name for the prior distribution and `a,b` with the prior's parameters. [Scipy docs](https://docs.scipy.org/doc/scipy/reference/stats.html) are your best friend! If any line of the stub below is unfamiliar/unclear you should look at the appropriate library documentation. We will need these kinds of libraries all the time, so you must be comfortable with them.


_Points:_ 1

```python
import matplotlib.pyplot as plt
import numpy as np
import seaborn as sns
%matplotlib inline
%config InlineBackend.figure_format='retina'
sns.set_style('white')
xvs = np.arange(0,1,0.01)

# prior parameters
DISTRIBUTION = ...
a = ...
b = ...

# the statement below will import the function you defined as a string above
exec(f'from scipy.stats import {DISTRIBUTION}')

def calc_pdf(xvs, a, b):
    """
    for each x value in the xvs, calculate their 
    corresponding pdf of the distribution of your choice

    Return: a list of pdf that has the same length as x_axis
    """
    ...

# likelihood parameters from the data 
n = ...
x = ...

# put your MLE solution from Q7.5 here
mle = ...


def plot_74():
    # for each y_values, use `calc_pdf` to map the x-axis to the y
    plt.plot(
        xvs, 
        ...
        label=f'prior a={a},b={b}' 
    )
    
    plt.plot(
        xvs, 
        ...
        label=f'posterior | data x={x},n={n}' 
    )
    
    plt.axvline(
        mle, 
        linestyle="--",
        color='r',
        label=f'MLE={mle}'
    )
    
    sns.despine()
    plt.legend(loc='upper left');
    plt.title('Uninformative prior');

plot_74()
```

```python
grader.check("74_bayesian_est")
```

<!-- BEGIN QUESTION -->

### Q7.5 Graph Interpretation

Describe what you see in the previous graph. Which distribution the uninformative prior looks like? How does prior influence the posterior? How does the MLE similar or different from the posterior distribution with this particular prior?

_Points:_ 0.2

```python
plot_74()
```

_Type your answer here, replacing this text._

<!-- END QUESTION -->

## Q8 Maximum A Posteriori (MAP) (2 Points)

In Q8 we were doing full Bayesian estimation... the plots of the posterior in the questions above show the distribution of possible parameters. But sometimes we are interested in where the maximum value of the posterior distribution lies; this peak is called the Maximum A Posteriori estimate. 

<!-- BEGIN QUESTION -->

### Q8.1 Intuition of MAP

What are the similarities between MLE and MAP? What does MAP take into account that MLE does not? 

What are the similarities between MAP and Bayesian estimation? What does Bayesian take into account that MAP does not?


_Points:_ 0.4

_Type your answer here, replacing this text._

<!-- END QUESTION -->

<!-- BEGIN QUESTION -->

### Q8.2 MAP estimate for the coin flipping dataset 

Recall that the posterior of the coin flipping task is the same kind of distribution as the prior.

What kind of distribution is the posterior? What is the generic equation for the peak (i.e., mode) of that distribution? In words, describe why that is the peak.

Write down the equations for the MAP estimate of the coin flipping task. Parameterize it in terms of the prior parameters `a,b` and the likelihood parameters `x` (the number of heads) and `n` (the number of total flips). It may help you to look at the answer to Q7.2

_Points:_ 0.4

_Type your answer here, replacing this text._

<!-- END QUESTION -->

### Q8.3 Plotting MAP
For this version of the coins example... lets say that the prior should be a symmetric prior centered at $\theta=0.5$ that indicates we think this coin is pretty likely to be fair. Pick any set of prior parameters $a,b$ that reflect this requirement.

As before the observed data are (0=tails,1=heads): $\mathcal{D}_1 = [0, 1, 1, 1, 1, 1, 0, 1, 1, 1]$

Use the code below and the parameters above to plot MLE, the full Bayesian estimate of the posterior distribution, and the MAP.


_Points:_ 0.8

```python
import matplotlib.pyplot as plt
import numpy as np
import seaborn as sns
%matplotlib inline
sns.set_style('white')
xvs = np.arange(0,1,0.005)

# prior parameters
DISTRIBUTION = ...

# careful about the name. Use a_1 and b_1 instead of a and b
a_1 = ...
b_1 = ...

# the statement below will import the function you defined as a string above
exec(f'from scipy.stats import {DISTRIBUTION}')

# likelihood parameters from the data 
n = ...
x = ...

# put your MLE solution from Q7.5 here
mle = ...

# put your MAP solution from Q8.2 here
# careful about the name. Use a_1 and b_1 instead of a and b
MAP = ...

def plot_83():
    # careful about the name. Use a_1 and b_1 instead of a and b
    # for each y_values, use `calc_pdf` from Q8 to map the x-axis to the y
    plt.plot(
        xvs, 
        ...
        label=f'prior a={a_1},b={b_1}' 
    )
    
    plt.plot(
        xvs, 
        ...
        label=f'posterior | data x={x},n={n}' 
    )
    
    # plot the MLE line
    plt.axvline(
        mle, 
        color='r',
        linestyle="--",
        label=f'MLE={mle}'
    )
    
    # plot the MAP line
    plt.axvline(
        MAP, 
        color='g',
        linestyle="--",
        label=f'MAP={MAP}'
    )
    
    sns.despine()
    plt.legend(loc='upper left');
    plt.title('Previous data prior');

plot_83()
```

```python
grader.check("83_MAP")
```

<!-- BEGIN QUESTION -->

### Q8.4 Graph Interpretation

Describe what you see in the previous graph. How does prior influence the posterior? Compare the MLE and MAP estimate. Explain why the two estimate value are different, or the same.

_Points:_ 0.4

```python
plot_83()
```

_Type your answer here, replacing this text._

<!-- END QUESTION -->

## The End of A1!

This is the end of the A1.

Have a look back over your answers, and also make sure to `Restart & Run All` from the kernel menu to double check that everything is working properly. This restarts everything and runs your code from top to bottom.

Once you're happy with your work, click the disk icon to save, and submit the zip file onto gradescope. **You MUST submit all the required component to receive credit.**

Note that you can submit at any time, but **we grade your most recent submission**. This means that **if you submit an updated notebook after the submission deadline, it will be marked as late**.

_Note:_ If you encountered `LatexFailed` message during exporting, or has the `AttributeError: module 'nbconvert' has no attribute 'pdf'` error, it indicates that your $\LaTeX$ code is not correct. Try to the LaTeX syntax error by scrolling up to see the LaTeX error message. If you need any additional helps, please make a private post on campuswire and we are happy to help.

## Submission

Make sure you have run all cells in your notebook in order before running the cell below, so that all images/graphs appear in the output. The cell below will generate a zip file for you to submit.

Please make sure to see the output of the gradescope autograder. You are responsible for waiting and ensuring that the autograder is executing normally for your submission. Please create a campuswire post if you see errors in autograder execution.

```python
grader.export(force_save=True, run_tests=True, files=['imgs'])
```

Tests for 74_bayesian_est (points: 1)
```
>>> assert isinstance(a, int), 'is your a an integer'
>>> assert isinstance(b, int), 'is your b an integer'
>>> assert isinstance(x, int), 'is your x an integer'
>>> assert isinstance(n, int), 'is your n an integer'
>>> assert isinstance(mle, float), 'is your mle a float'
>>> assert isinstance(DISTRIBUTION, str), 'is your prior/posterior a string describing a scipy.stats distribution'
```


Tests for 83_MAP (points: 0.8)
```
>>> assert isinstance(a_1, int), 'is your a an integer'
>>> assert isinstance(b_1, int), 'is your b an integer'
>>> assert isinstance(x, int), 'is your x an integer'
>>> assert isinstance(n, int), 'is your n an integer'
>>> assert isinstance(mle, float), 'is your mle a float'
>>> assert isinstance(DISTRIBUTION, str), 'is your prior/posterior a string describing a scipy.stats distribution'
```
